Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rebecca Mqamelo"
COLLABORATORS = "None"
```

---

# CS110 Pre-class Work 1.2



## Question 1
The following pseudocode has been extracted from Cormen et al. You should analyze the lines of the pseudocode trying to understand what every line in the pseudocode does and how many times is that line executed. 
![Insertion sort pseudo-code](insertionsort_pseudo.png)

Answer the following questions: What does line 7 do? Why is it necessary?

<!-- grade_id: cell-3bac59ff24b11a27, points: 0, grade: true, solution: true -->
The item just before the index being sorted has moved forward one place, with j (the key) replacing that index. Since i is references as always being one index behind j, i must be updated after the sort has executed to reflect "moving back" with j.

## Question 2
The following Python code is an attempt to implement the Insertion-Sort pseudocode above. However, this code has a bug. 


```python
def insertionSort(A):
    for j in range(len(A)):
        key = A[j]
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
            A[i+1] = key
    return A
```

Copy and paste the code in the cell below. Identify and correct the bug. 

<!-- grade_id: cell-1b7488e89cb501c5, points: 0, grade: false, solution: true -->
```python
def insertionSort(A):
    for j in range(len(A)): # Termination condition
        key = A[j] # Initialization:
                   # Loop invariant holds before the first iteration
        i= j-1
        while i >= 0 and A[i]>key:
            A[i+1] = A[i]
            i -= 1
        A[i] = key # Maintenance:
                   # Incrementing j for the next iteration
                   # of the for loop preserves the 
                   # loop invariant.
    return A
    raise NotImplementedError()
```

<!-- grade_id: cell-3432ce8ae2493fa6, points: 1, grade: true, solution: false -->
```python
assert(insertionSort([0]) == [0])
assert(insertionSort([-1,1]) == [-1,1])
assert(insertionSort([1,-1]) == [-1,1])
assert(insertionSort([1,6,3,6]) == [1,3,6,6])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-15-1e36be01afe1> in <module>
      1 assert(insertionSort([0]) == [0])
----> 2 assert(insertionSort([-1,1]) == [-1,1])
      3 assert(insertionSort([1,-1]) == [-1,1])
      4 assert(insertionSort([1,6,3,6]) == [1,3,6,6])

AssertionError:
```

How do you know this code works as intended now?

<!-- grade_id: cell-442998d5fdd8b561, points: 0, grade: true, solution: true -->
The assert method should raise no errors i.e. produce a blank output

## Question 3 
Implement the bubble sort algorithm in the cell below. First, look up the corresponding pseudocode in Cormen et al. (page 40), and write a Python implementation for the bubble sort.  Like in the `insertionSort`,  your new Python function, `bubbleSort`, takes a list of elements as an argument and returns the sorted list.

<!-- grade_id: cell-b709733c96d8b615, points: 0, grade: false, solution: true -->
```python
def bubbleSort(A):
    for i in range(len(A)-1):
        for j in range(i+1, len(A)): # Why are we ignoring the first index?
            if A[j] < A[j-1]:
                A[j] = A[j-1] 
    raise NotImplementedError()
```

<!-- grade_id: cell-ffc9b251ea0047ba, points: 2, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

Is your program bug-free? Think of three different [“test cases”](https://www.hiredintech.com/classrooms/algorithm-design/lesson/83) to check the correctness of your code. If a code passes the tests you have just outlined, is this an unambiguous proof that the code is correct? Give your answer in the cell below. Feel free to use additional cells (Markdown or Code) as you wish.

<!-- grade_id: cell-5ecb51404140cb2f, points: 0, grade: true, solution: true -->
No, it looks like it isn't bug free...

```python
assert(bubbleSort([0]) == [0])
assert(bubbleSort([-1,1]) == [-1,1])
assert(bubbleSort([1,-1]) == [-1,1])
assert(bubbleSort([1,6,3,6]) == [1,3,6,6])
```

Output:
```
<ipython-input-10-3581d7cb95bc> in <module>
----> 1 assert(bubbleSort([0]) == [0])
      2 assert(bubbleSort([-1,1]) == [-1,1])
      3 assert(bubbleSort([1,-1]) == [-1,1])
      4 assert(bubbleSort([1,6,3,6]) == [1,3,6,6])

<ipython-input-9-44bee2a8f34c> in bubbleSort(A)
      5                 A[j] = A[j-1]
      6 
----> 7     raise NotImplementedError()

NotImplementedError:
```

## Question 4
The following pseudocode corresponds to a third sorting algorithm to be covered in Lesson 1.2, the selection sort.  Like in the previous step, write a Python implementation of the algorithm and make sure that your program is bug-free. `selectionSort` should take a list of elements as an argument and return the sorted list.

![Selection sort pseudo-code](selectionsort_pseudo.png)

<!-- grade_id: cell-e04f94e5d418923c, points: 0, grade: false, solution: true -->
```python
def selectionSort(A):
    n = len(A)
    for i in range(n-1):
        minidx = i
        for j in (i+1, n):
            if A[j] < A[minidx]:
                minidx = j
            A[i] = A[minidx] 
    raise NotImplementedError()
```

<!-- grade_id: cell-ffe912ac7922de70, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

```python
assert(selectionSort([0]) == [0])
assert(selectionSort([-1,1]) == [-1,1])
assert(selectionSort([1,-1]) == [-1,1])
assert(selectionSort([1,6,3,6]) == [1,3,6,6])
```

Output:
```
<ipython-input-13-ae060ba1cca3> in <module>
----> 1 assert(selectionSort([0]) == [0])
      2 assert(selectionSort([-1,1]) == [-1,1])
      3 assert(selectionSort([1,-1]) == [-1,1])
      4 assert(selectionSort([1,6,3,6]) == [1,3,6,6])

<ipython-input-12-f90f4df3df36> in selectionSort(A)
      7                 minidx = j
      8             A[i] = A[minidx]
----> 9     raise NotImplementedError()

NotImplementedError:
```

**Please make sure you try your best to code insertion sort, selection sort and bubble sort, and bring these programs to class. You will need them for a class activity.**