Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rebecca"
COLLABORATORS = "None"
```

---

# CS110 Pre-class Work 2.2

## Question 1 (Exercise 3.1-3 of Cormen, et al. )
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
Big O notation does not tell us anything about the time or space and algorithm takes. Rather, it is an upper bound assymptote. It describes an asymptotic function because "For large enough inputs, the multiplicative constants and lower-order terms of an exact running time are dominated by the effects of the input size itself."

## Question 2 (Exercise 3.1-4 of Cormen, et al. )

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
1. $2^{n+1}=O(2^n)$ because $2^{n+1}=2(2^{n})=O(2n)$
2. No. Because $2^{2n}=(2^{n})^{n}$ which does not equal $O(2^n)$

## Question 3.
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
import math
from math import inf
    
def max_subarray_recursion(A, low, high):
    if high == low:
        return(low, high, A[low]) # base case = only one element
    else: 
        mid = int(math.floor(float(low+high)/2))
        # recursive case:
        left_low, left_high, left_sum = max_subarray_recursion(A, low, mid)
        right_low, right_high, right_sum = max_subarray_recursion(A, mid+1, high)
        cross_low, cross_high, cross_sum = max_crossing_subarray(A, low, mid, high)
        
        final = max(left_sum,right_sum,cross_sum)        
        if left_sum == final:
            return left_low,left_high,left_sum
        elif right_sum == final:
            return right_low,right_high,right_sum
        else:
            return cross_low,cross_high,cross_sum
     
        # comparison system I had before:
        # if left_sum >= right_sum and left_sum >= cross_sum:
            # return(left_low, left_high, left_sum)
        # elif right_sum >= left_sum and right_sum >= cross_sum:
            # return(right_low, right_high, right_sum)
        # else:
            # return(cross_low, cross_high, cross_sum) 

def max_crossing_subarray(sublist, low, mid, high):
    inf = float('inf')
    left_sum = -inf
    sum = 0
    i = mid
    while i >= low:
        sum = sum + sublist[i]
        if sum > left_sum:
            left_sum = sum
            left_index = i
        i-= 1 # moving left from the midpoint
    
    right_sum = -inf
    sum = 0
    j = mid+1
    while j <= high:
        sum = sum + sublist[j]
        if sum > right_sum:
            right_sum = sum
            right_index = j
        j+= 1 # moving right from the midpoint
        
        
    for i in range(mid, low, -1):
        sum = sum + A[i]
        if sum > left_sum:
            left_sum = sum
            max_left = i
    
    total = left_sum + right_sum
    return(left_index, right_index, total)
   
    j = mid + 1

max_subarray_recursion([-2,1,-1,2,-5], 0, 4)
max_subarray_recursion([-2, -5, 6, -2, -3, 1, 5, -6], 0, 7)
```

Output:
```
<ipython-input-5-fa324c73d6a4> in max_subarray_recursion(A, low, high)
     10         left_low, left_high, left_sum = max_subarray_recursion(A, low, mid)
     11         right_low, right_high, right_sum = max_subarray_recursion(A, mid+1, high)
---> 12         cross_low, cross_high, cross_sum = max_crossing_subarray(A, low, mid, high)
     13 
     14         final = max(left_sum,right_sum,cross_sum)

<ipython-input-5-fa324c73d6a4> in max_crossing_subarray(sublist, low, mid, high)
     52 
     53     for i in range(mid, low, -1):
---> 54         sum = sum + A[i]
     55         if sum > left_sum:
     56             left_sum = sum

IndexError: list index out of range
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

Output:
```
<ipython-input-5-5ddab6f2f574> in <module>
----> 1 assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
      2 assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))

<ipython-input-4-84621ab31064> in bruteforce_max_subarray(A)
     11     """
     12     # YOUR CODE HERE
---> 13     raise NotImplementedError()

NotImplementedError:
```

## Question 4. 
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7]
max_subarray_recursion(A, 0, len(A)-1)
```

Output:
```
(7, 10, 59)
```

## Question 5. Asymptotic notation. 
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  |
Selection sort |  |
Bubble sort |  | 
Finding maximum subarray |  |

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->
Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort | $On^{2}$ | More correct: $\Theta n^{2}$
Selection sort | $On^{2}$ | More correct: $\Theta n^{2}$
Bubble sort    | $On^{2}$ | More correct: $\Theta n^{2}$
Finding maximum subarray |  | $\Theta n(log n)$

## [Optional] Question 6. 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```