Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rebecca"
COLLABORATORS = "None"
```

---

# CS110 Pre-class Work 11.2

## Part A.  Currency trading (Slightly simplified version of 15.3-6 from Cormen et al.)

Imagine that you wish to exchange one currency for another. You realize that instead of directly exchanging one currency for another, you might be better off making a series of trades through other currencies, winding up with the currency you want.

Suppose that you can trade n different currencies, numbered $1,2,… ,n$, where you start with currency 1 and wish to wind up with currency $n$. You are given, for each pair of currencies $i$ and $j$ , an exchange rate $r_{ij}$ , meaning that if you start with $d$ units of currency $i$ , you can trade for $dr_{ij}$ units of currency $j$. Note that the total number of trades allowed is limited to $n$.

Assuming there is no commission, propose an algorithm (described in prose) to solve this problem using either a memoization or bottom-up strategy. **Note that we will be working on the Python implementation in class.**

Click [here](https://drive.google.com/open?id=1L8Fjo1Xt8sltab-tz3m91eTuiE9LYWF7) for some example data.

<!-- grade_id: cell-3f8a8808c720baf2, points: 0, grade: true, solution: true -->
This is an optimization problem similar tyo the rod-cutting scenario, where we want whatever combination of currency trades we choose to leave us with the most vlaue in the end (i.e. we want the least positive arbitrage possible on our trades – ideally, since this a zero-commission example, we might even end up with more value (money) than we started with.) We can use dynamic programming using an array in which we store the result of each exchange trade, perhaps in the units of the original currency. Once we've solved this for one direction in ine currency pair, we store that as a unique trade, which might be called upon later if we choose to start with exchanging with another currency. In other words, we want an array that will store every possible combination ($2^n$) of currency exchanges. This calls for a matrix, which is exactly what ynamic programming provides.

## Part B - Money game (Solution)
Consider a row of n coins of values $v_1, v_2,...,v_n$, **where $n$ is even**. We play a game against an opponent by alternating turns. In each turn, a player selects either the first or last coin from the row, removes it from the row permanently, and receives the value of the coin. Determine the maximum possible amount of money we can definitely win if we move first.

For example, consider the game:

\$2, \$10, \$1, \$5

By moving first and playing optimally one can be guaranteed of \$15. The first move is to take \$5. This forces your opponent to take either \$2 or \$1, and then allows you to take \$10. Assume that the opposing player also plays optimally (i.e., minimizing the gain of the first-move player.)

As a hint, we will provide the recurrence for the solution to this problem. Suppose $c(i,j)$ is the maximum possible amount of money the first player can win for the row of coins of value $v_i, v){i+1},...,v_j$ (so to solve our problem, we need to compute $c(1,n)$. Then:

$$c(i,j)=max[v_i+min(c(i+2,j), c(i+1,j-1)), v_j+min(c(i,j-2), c(i+1,j-1))]$$

You are also encouraged to watch [this video](https://www.youtube.com/watch?v=KnP8_L13xW4&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5) if you need some more hints.

## Question 1. 
Complete the following function to solve the game using a bottom-up strategy, assuming that the opposing player also plays optimally. 


<!-- grade_id: cell-4014ae901b66304b, points: 0, grade: false, solution: true -->
```python
def bottom_up_coin_game(S):
    """
    Returns the maximum possible amount of money the first-move player can win,
    given an array of coin values. The function solves this using a bottom-up 
    approach.
    
    Inputs:
    - S: list of floats, values of the coins, of even length. 
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. max_val is None when the length of list A is odd.
    """
    # We need n+1 rows as the table is constructed  
    # in bottom up manner using the base case 0 value 
    # case (n = 0) 
    
    m = 0
    n = len(S)-1
    
    table = [[0 for x in range(m)] for x in range(n+1)] 
  
    # Fill the entries for 0 value case (n = 0) 
    for i in range(m): 
        table[0][i] = 1
  
    # Fill rest of the table entries in bottom up manner 
    for i in range(1, n+1): 
        for j in range(m): 
  
            # Count of solutions including S[j] 
            x = table[i - S[j]][j] if i-S[j] >= 0 else 0
  
            # Count of solutions excluding S[j] 
            y = table[i][j-1] if j >= 1 else 0
  
            # total count 
            table[i][j] = x + y 
  
    max_val = table[n][m-1]
    return  max_val
  
# Driver program to test above function 
arr = [1, 2, 3] 
m = len(arr) 
n = 4
print(bottom_up_coin_game(arr))
```

Output:
```
<ipython-input-8-87165a7f3498> in <module>
     45 m = len(arr)
     46 n = 4
---> 47 print(bottom_up_coin_game(arr))

<ipython-input-8-87165a7f3498> in bottom_up_coin_game(S)
     38             table[i][j] = x + y
     39 
---> 40     max_val = table[n][m-1]
     41     return  max_val
     42 

IndexError: list index out of range
```

```python
import numpy as np
assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)  
assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)  
assert(bottom_up_coin_game([]) == 0)  
```

Output:
```
<ipython-input-9-b7e49a2da0fa> in <module>
      1 import numpy as np
----> 2 assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)
      3 assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)
      4 assert(bottom_up_coin_game([]) == 0)

<ipython-input-8-87165a7f3498> in bottom_up_coin_game(S)
     38             table[i][j] = x + y
     39 
---> 40     max_val = table[n][m-1]
     41     return  max_val
     42 

IndexError: list index out of range
```

<!-- grade_id: cell-f7a5d6c26e95dab2, points: 1, grade: true, solution: false -->
```python
import numpy as np
assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)  
assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)  
assert(bottom_up_coin_game([]) == 0)  
```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-5-b7e49a2da0fa> in <module>
      1 import numpy as np
----> 2 assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)
      3 assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)
      4 assert(bottom_up_coin_game([]) == 0)

TypeError: bottom_up_coin_game() missing 2 required positional arguments: 'm' and 'n'
```

## [Optional]  Question 2. 

Complete the function `print_strategy` to print out a solution to the coin game. Completing the helper function `extended_bottom_up_coin_game` to use in `print_strategy` is optional.

<!-- grade_id: cell-4583755bc245e1e9, points: 0, grade: false, solution: true -->
```python
def print_strategy(A):
    """
    Print coins to pick for the first player
    
    Inputs:
    - A: list of floats, values of the coins, of even length.
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. 
        * max_val is None if the length of list A is odd.
        * max_val is 0 if list A is empty (no coins)
    - out: list of values of the coins that the first player picks in that
    order.
        * out is None if max_val is 0 or None.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-75ce551c5185e374, points: 0, grade: true, solution: false -->
```python
assert(print_strategy([4,6,5,20,4,6,4,8])[0] == 40)
assert(print_strategy([4,6,5,20,4,6,4,8])[1] == [8, 6, 6, 20])
```