```python
import math
import random
import numpy as np
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rebecca"
COLLABORATORS = "None"
```

---

# CS110 Pre-class Work 4.2

## Part A. The Hire-Assistant Problem.

Imagine that you need to hire a new assistant. Every day an agency sends a new assistant for you to interview. If the assistant is better than your current assistant, then you fire your current assistant and you hire the better assistant. You may assume that assistant quality is uniformly distributed between 0 and 1.

## Question 1.
Write a function, named hire_assistant, that takes applicants (a list of the numbers that represent the level of qualification of the applicants; the higher the number, the better qualified), and returns the number hires if the applicants are presented in the exact same order as the input list applicants. Note that your function should not randomize anything (or else it would be called a randomized algorithm).

<!-- grade_id: cell-d011f5f4707fe41a, points: 0, grade: false, solution: true -->
```python
def hire_assistant(applicants):
    """
    Return the number of assistant hired.
    Inputs:
    - applicants: a list of the numbers that represent the level of qualification of 
    the applicants; the higher the number, the better qualified
    
    Outputs:
    - hires: Number of assistants hired
    """
    # YOUR CODE HERE
    # We assume no one is currently employed, therefore the best candidate is the first candidate we see.
    best = applicants[0]  
    hired = []
    for i in applicants:
        interview_candidate = i
        if interview_candidate > best: 
            # Note: I did not make it inclusive (>=)
            # because it doesn't make sense to incur the cost of hiring someone
            # if they aren't any better than the person you currently have.
            # This is why my function will not satisfy "assert" below.
            best = interview_candidate
            print("Congratulations! You just made a new hire!")
            hired.append(best)
        else:
            print("No hire this time...")
    return len(hired) 

hire_assistant([1])
print('')
hire_assistant([-1, -2, -3, -4])
print('')
hire_assistant([1, 2, -2, 4])
print('')
hire_assistant([1, 2, 3, -6, 8])
```

Output:
```
No hire this time...

No hire this time...
No hire this time...
No hire this time...
No hire this time...

No hire this time...
Congratulations! You just made a new hire!
No hire this time...
Congratulations! You just made a new hire!

No hire this time...
Congratulations! You just made a new hire!
Congratulations! You just made a new hire!
No hire this time...
Congratulations! You just made a new hire!
3
```

<!-- grade_id: cell-66778b97ad66f71e, points: 1, grade: true, solution: false -->
```python
assert(hire_assistant([1])==1)
assert(hire_assistant([-1, -2, -3, -4])==1)
```

Output:
```
No hire this time...
AssertionError                            Traceback (most recent call last)
<ipython-input-5-2dddc1aba31f> in <module>
----> 1 assert(hire_assistant([1])==1)
      2 assert(hire_assistant([-1, -2, -3, -4])==1)

AssertionError:
```

## Question 2. 
Assuming the applicants are presented in a random order, write a function that receives the number of applicants as input and returns the average number of assistants hired.

**N.B.:** Don’t forget to run the simulation several times for each given number of applicants to better estimate the number of hires (please refer to task 3 of the Study Guide).

<!-- grade_id: cell-b223520ca72942a0, points: 0, grade: true, solution: true -->
```python
import random
def experimental_hires(N):
    # YOUR CODE HERE
    hired = []
    best = applicants[0]
    # Generate a random order of applicants with random qualities
    applicants = np.random.random(N)
    for interview_candidate in applicants:
        if interview_candidate > best:
            best = interview_candidate
            hired.append(best)
    return len(hired)
```

## Question 3.

Use the function below, `analytical_hires(N)`, which returns the analytical expected number of hires, given the number of applicants, along with the function you created in question 2 to create a graph with two curves such that:
* The x-axis shows the total number of applicants (make sure label the x-axis)
* The y-axis shows the average number of hires (make sure label the y-axis)
* The graph contains two curves;
    * Curve 1: the theoretical performance estimates computed calls to the function `analytical_hires`.
    * Curve 2: the simulated or experimental estimates using the function you created in question 2.


```python
def analytical_hires(N):
    """
    Return the analytical expected number of hires if there are N applicants
    Inputs:
    - N: Number of applicants
    Outputs:
    - hires: Average number of assistants hired
    """
    # from the textbook, we know that the analytical result is 
    # 1 + 1/2 + 1/3 + ... + 1/N
    hires = 0
    for n in range(N):
        hires += 1/(n+1)
    return hires
```

<!-- grade_id: cell-f9c07920c069ce20, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def simulation(n,m):
    num_applicants = [i for i in range(n)]
    num_hired = []
    for i in range(n):
        count = 0
        for j in range(m):
            count += hire_assistant(i)
        average_i = count / m
        num_hired.append(average_i)
    plt.subplot(1,2)
    plt.plot(num_applicants, num_hired)
    plt.show
    return num_applicants, num_hired

simulation(1000,10)
```

Output:
```
<ipython-input-22-79062a8300c1> in simulation(n, m)
      6         count = 0
      7         for j in range(m):
----> 8             count += hire_assistant(i)
      9         average_i = count / m
     10         num_hired.append(average_i)

<ipython-input-4-f715a232301d> in hire_assistant(applicants)
     11     # YOUR CODE HERE
     12     # We assume no one is currently employed, therefore the best candidate is the first candidate we see.
---> 13     best = applicants[0]
     14     hired = []
     15     for i in applicants:

TypeError: 'int' object is not subscriptable
```

## Question 4.

Plot a graph with the x-axis showing the total number of applicants and the y-axis showing the probability that exactly one assistant is hired.

<!-- grade_id: cell-d3fe1b7d6d175ad7, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-15-15b94d1fa268> in <module>
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

## [Optional] Question 5.
Assume that an assistant is able to perform an amount of work each day that is equal to their “quality”. You have a total amount of work M that needs to be accomplished. Your costs are as follows:
* X = daily salary for the assistant,
* Y = fee to the employment agency,
* Z = retrenchment fee for the old assistant.

Try to formulate an optimal stopping rule (ie. at what point should one stop requesting new potential hires from the agency?) Make any necessary assumptions to ensure the problem is well-formulated.


<!-- grade_id: cell-af2f0291eced6982, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-16-15b94d1fa268> in <module>
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

## Part B. The Hat Check Problem.

There is a coat check at a party, where an attendant stores everyone’s hat while they attend the party. The attendant receives the N hats from everyone attending (all attendees come with a hat). Unfortunately, the coat check attendant forgets which hat belongs to whom. Rather than admitting a mistake, the attendant simply returns random hats back to the party goers. 
What is the average number of correct hats returned? Here are some guiding questions to help you to simulate this problem. 

## Question 1. 
Knowing that everyone’s hats are unique and every guest has a hat. Do you need to generate a random sample in a similar way as what you did for the hiring assistant problem? 

<!-- grade_id: cell-e786799fc4eb1499, points: 0, grade: true, solution: true -->
Yes, because we are purposefully returning random hats.

## Question 2. 
Which of the following commands do you think is the Pythonic way to implement that? 
```
import numpy as np
n = 100 #the number of party attendants `
```
**Command 1. **
```
hat_list = [np.random.integers(0,n) for i in range(n)]`
```
**Command 2.**
```
hat_list = list(range(n)) 
np.random.shuffle(hat_list) 
```
**Command 3.**
```
hat_list = np.random.sample(n)
```

<!-- grade_id: cell-b8da78e72c1c0738, points: 0, grade: true, solution: true -->
Command 2

## Question 3.
Now write a function `hat_check(N)` that has: 
* Input: N the number of party attendants. 
* Output: the number of hats correctly returned despite the fact that hats are randomly handed back to the guests.

You should use the command you picked for question 2. 

<!-- grade_id: cell-c8499aeb1b1d76c7, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def hat_check(n):
    # original hats and people
    people = [i for i in range(0,n)]
    # hats given to people later on
    hats = random.sample(range(0,n),n)
#     print(hats)
    matched = 0
    # check if there is any match
    for i in range(n):
        if people[i] == hats[i]:
            matched += 1
    return matched
```

## Question 4.

Plot a curve with the x-axis showing the total number of party attendants and the y-axis showing the average number of hats correctly returned. As always, remember to run several trials. 

<!-- grade_id: cell-597031ea2a5a512a, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def simulation(n,m):
    num_people = [i for i in range(n)]
    num_matched = []
    for i in range(n):
        count = 0
        for j in range(m):
            count += hat_check(i)
        average_i = count / m
        num_matched.append(average_i)

    plt.subplot(1,2)
    plt.plot(num_people, num_matched)
    plt.show
    print(num_people, num_matched)
    
simulation(100,10)
```

Output:
```
//anaconda3/lib/python3.7/site-packages/matplotlib/figure.py in add_subplot(self, *args, **kwargs)
   1412                     self._axstack.remove(ax)
   1413 
-> 1414             a = subplot_class_factory(projection_class)(self, *args, **kwargs)
   1415 
   1416         return self._add_axes_internal(key, a)

//anaconda3/lib/python3.7/site-packages/matplotlib/axes/_subplots.py in __init__(self, fig, *args, **kwargs)
     62                 # num - 1 for converting from MATLAB to python indexing
     63         else:
---> 64             raise ValueError(f'Illegal argument(s) to subplot: {args}')
     65 
     66         self.update_params()

ValueError: Illegal argument(s) to subplot: (1, 2)
<Figure size 432x288 with 0 Axes>
```

## [Optional] Question 5.
As $N$ tends to infinity, the number of correct hats returned tends towards a well-known statistical distribution. State the distribution with all its parameters. Plot several samples using your code. Does the empirical distribution match your theoretical prediction?

<!-- grade_id: cell-32fe26c1d99fdd2a, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE