Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rebecca"
COLLABORATORS = "None"
```

---

# CS110 Pre-class Work 9.1

## Part A. Average number of comparisons when searching

## Question 1. 

Complete the following function that computes the depth of a given node in a BST identified by its root:

```python
import random
```

<!-- grade_id: cell-cd10e9a432673bd9, points: 0, grade: false, solution: true -->
```python
def depth(root, node):
    """finds the depth of the node in a BST. depth of root is 0.

    Input:
    - root: a node, the root of the BST
    - node: a node to compute the depth of

    Output:
    - d: int, distance from node to root"""
    if root is None:
        return 0
    else:
        # Find height of nodes above thise one + 1 = node depth
        d = max(BST.height(root.left), BST.height(root.right)) + 1
        return d
```

<!-- grade_id: cell-c57bc5138040968c, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 2.

Complete the following function, making use of `depth`, to calculate the average number of comparisons required to search for a randomly chosen element of a standard BST. For example, the following tree:

                            12
                              \
                               14

will have an average number of comparisons of 1.5. This is because 50% of the time we will be searching for 12 (1 comparison), and 50% of the time we will be searching for 14 (2 comparisons).

You can test your function with this test case by constructing the tree using the code from the last session's pre-class work. 


(*Hint*: the number of comparisons required to search for node a is `depth(root, a) + 1`.) 

<!-- grade_id: cell-34e9cb302f1dad00, points: 0, grade: false, solution: true -->
```python
def depth_of_cmp(bst, depth, lst):
    if bst.l_child is not None:
        lst = depth_of_cmp(bst.l_child, depth+1, lst)
    if bst.r_child is not None:
        lst = depth_of_cmp(bst.r_child, depth + 1, lst)
    lst.append(depth)
    return lst

def avg_cmp(bst):
    """Finds the average number of comparisons required 
    to search for a randomly chosen element of a standard BST. 
    
    Input:
    - bst: a Node, the root of the BST
    
    Output:
    - n: int, average number of comparisons"""
    
    lst = []
    depth_of_cmp(bst, 0, lst)
    return np.mean(lst)
```

<!-- grade_id: cell-dedc229bbb15d811, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Part B. Depth of randomly-built trees
The average number of comparisons for a randomly chosen element within the tree is related to the “average depth” of the tree. Here the average depth of a tree is defined to be the average of the depths of all the nodes in the tree. 
## Question 1. 

Complete the following function to find the maximum depth of a node in a given BST. The function must run in $\mathrm{O}(N)$ time and use $\mathrm{O}(h)$ space (where $N$ is the number of elements in the BST and $h$ is the height of the tree). Also, note that the maximum depth is equal to the height of the tree.


<!-- grade_id: cell-64aba253c2b1bada, points: 0, grade: false, solution: true -->
```python
def max_depth(bst, depth=0, max_height=0):
    """Finds the maximum depth of node in a BST.
    
    Input:
    - bst: a Node, the root of the BST
    
    Output:
    - h: int, the maximum depth in a BST"""
    
    if bst.l_child is not None:
        max_height = max_depth(bst.l_child, depth+1, max_height)
    if bst.r_child is not None:
        max_height = max_depth(bst.r_child, depth+1, max_height)
    max_height = max(depth, max_height)
    return max_height
```

<!-- grade_id: cell-82deabcb0ea2e594, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 2. 

Using the `avg_cmp` function, complete the following function to return the average depth of the tree. The average depth is related to the average number of comparisons as `average depth = (average comparisons - 1)`, since the root has depth zero.

<!-- grade_id: cell-347916904e4a8379, points: 0, grade: false, solution: true -->
```python
def avg_depth(bst):
    """Computes the average depth of a BST
    
    Input:
    - bst: a Node, root of the BST
    
    Output:
    - avg_d: float, average depth of the BST"""
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1bd7535fe6950d5f, points: 1, grade: true, solution: false -->
```python
class Node:
    def __init__(self, val):
        self.l_child = None
        self.r_child = None
        self.parent = None
        self.data = val

    def insert(self, node):
        """inserts a node into a *non-empty* tree rooted at the node, returns
        the root"""
        if self.data > node.data:
            if self.l_child is None:
                self.l_child = node
                node.parent = self
            else:
                self.l_child.insert(node)
        else:
            if self.r_child is None:
                self.r_child = node
                node.parent = self
            else:
                self.r_child.insert(node)
        return self
    
    def minimum(self):
        node = self
        while node.l_child != None:
            node = node.l_child
        return node

    def search_data(self, value):
        """searches a *non-empty* tree rooted at the node for a node with
        data = value, returns the value if found, None otherwise"""
        node = self.search(value)
        if node:
            return node.data
        else:
            return node
        
    def search(self, value):
        """searches a *non-empty* tree rooted at the node for a node with data = value, 
        returns the node if found, None otherwise"""
        if value == self.data:
            return self
        if value < self.data:
            if self.l_child:
                return self.l_child.search(value)
            else:
                return None
        else:
            if self.r_child:
                return self.r_child.search(value)
            else:
                return None
    
    def delete(self, value):
        def transplant(node1, node2):
            """
            replaces the subtree rooted at node1 with the subtree rooted at node2
            """
            nonlocal root
            if node1.parent == None:
                root = node2
            elif node1 == node1.parent.l_child:
                node1.parent.l_child = node2
            else:
                node1.parent.r_child = node2
            if node2 != None:
                node2.parent = node1.parent

        """if a node with data = value is present in the tree rooted at Node, deletes that node and returns the root"""
        root = self
        node = root.search(value)
        if node:
            if node.l_child is None:
                transplant(node, node.r_child)
            elif node.r_child == None:
                transplant(node, node.l_child)
            else:
                y = node.r_child.minimum()
                if y.parent != node:
                    transplant(y, y.r_child)
                    y.r_child = node.r_child
                    y.r_child.parent = y
                transplant(node, y)
                y.l_child = node.l_child
                y.l_child.parent = y
        return root
    
    def inorder(self): 
        """returns a list of all data in the tree rooted at root produced using an in order traversal"""
        lst = []
        root = self
        def sub_inorder(root):
            if not root is None:
                sub_inorder(root.l_child)
                lst.append(root.data)
                sub_inorder(root.r_child)
        sub_inorder(root)
        return lst
    
bst = None
nodes = [Node(15), Node(6), Node(18), Node(3), Node(7), 
         Node(17), Node(20), Node(2), Node(4)]
for node in nodes:
    if not bst:
        bst = node
    else:
        bst.insert(node)
def is_equal(a,b):
    return abs(a-b) < 1e-5

assert(avg_depth(nodes[0]) == 16/9)
assert(is_equal(avg_depth(nodes[1]),1.2))
assert(is_equal(avg_depth(nodes[2]), 2/3))
assert(is_equal(avg_depth(nodes[3]), 2/3))
assert(avg_depth(nodes[4]) == 0)
assert(avg_depth(nodes[5]) == 0)
assert(avg_depth(nodes[6]) == 0)
assert(avg_depth(nodes[7]) == 0)
assert(avg_depth(nodes[8]) == 0)
```

Output:
```
<ipython-input-10-32f6fe96b51d> in <module>
    111     return abs(a-b) < 1e-5
    112 
--> 113 assert(avg_depth(nodes[0]) == 16/9)
    114 assert(is_equal(avg_depth(nodes[1]),1.2))
    115 assert(is_equal(avg_depth(nodes[2]), 2/3))

<ipython-input-9-e7957ff664fa> in avg_depth(bst)
      8     - avg_d: float, average depth of the BST"""
      9     # YOUR CODE HERE
---> 10     raise NotImplementedError()

NotImplementedError:
```

## Question 3.
Now, insert randomly shuffled lists into BSTs, and measure the average depth and the maximum depth. How do these statistics scale as you increase $N$, the number of nodes? Make sure to give a compelling argument that motivates this scaling behaviour (you do not need to provide a technical derivation).

<!-- grade_id: cell-bad918a333cabe43, points: 0, grade: true, solution: true -->
```python
# Test case
r = None
for i in range(100):
    r = insert(r, Node(int(100000*random.random())))
print avg_cmp(r)
print max_height(r)
```


## Question 4.

Produce a plot showing the scaling behavior that you saw of both the average depth and the maximum depth as a function of the length of the shuffled list. For a meaningful figure, be sure to scale out to a list of size 10,000 and average the timings 50 times. (You don’t have to sample each value from 1 to 10,000!). Is this plot in agreement with the theoretical result you obtained in the previous question? Explain. 


<!-- grade_id: cell-3034eb71070a6734, points: 0, grade: true, solution: true -->
```python
# Plot
lst_avg_cmp = []
lst_max_height = []
for i in range(500, 1001, 500):
    lst_ith_avg_cmp = []
    lst_ith_max_height = []
    for j in range(50):
        r = None
        for i in range(i):
            r = insert(r, Node(int(100000*random.random())))
        lst_ith_avg_cmp.append(avg_cmp(r))
        lst_ith_max_height.append(max_height(r))
    lst_avg_cmp.append(np.mean(lst_ith_avg_cmp))
    lst_max_height.append(np.mean(lst_ith_max_height))

plt.subplot(223)
plt.plot([x for x in range(500, 1001, 500)], lst_avg_cmp, label="average depth")
plt.plot([x for x in range(500, 1001, 500)], lst_max_height, label="maximum depth")
plt.xlabel("length")
plt.ylabel("depth")
plt.legend(bbox_to_anchor=(1.05, 1), loc=2, borderaxespad=0.)
plt.show()
```