Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rebecca"
COLLABORATORS = "None"
```

```python
import math, time, random
from math import ceil
import matplotlib.pyplot as plt
import numpy as np
import random
```

---

# CS110 Pre-class Work 5.2

## Question 1.
Using Figure 7.1 in Cormen et al. as a model, perform manually the partition process on the following list: A = [1,5,6,2,3,8,9,4,7]. You just need to specify the followings:
1. The array after the process is done.
2. The value of $i$ after the process is done.

<!-- grade_id: cell-7aa520f8af13679b, points: 0, grade: true, solution: true -->
A = [1,5,6,2,3,8,9,4,7]
1. q = "3" = A[4]
2. Split in half at q: [1,5,6,2,3] and [8,9,4,7]

## Question 2.
Code up a Python implementation of `partition(A, p, r)`, closely follow the pseudo-code in Cormen et al., p.172. Your function should return the index of the pivot in the array.

<!-- grade_id: cell-1ceb2600756c60ff, points: 0, grade: false, solution: true -->
```python
def partition(A,p,r):
    """
    Assume r<len(A) and p>=0
    """
    # Pivot element is the last element in the subarray:
    x = A[r] 
    # We compare the pivot to the first element in the subarray:
    i = p-1 
    for j in range(p,r):
        # If current element is smaller than pivot:
        if A[j] <= x: 
            # Move on to the next one...
            i+=1
            # But first swap the pivot with the element you just analysed
            A[i], A[j] = A[j], A[i]
    # Swap the next element with the pivot so it becomes the new pivot and the 
    A[i+1], A[r] = A[r], A[i+1] 
    return i+1
    
    raise NotImplementedError()
```

<!-- grade_id: cell-a57b60117a7b82fb, points: 1, grade: true, solution: false -->
```python
A = [1,5,6,2,3,8,9,4,7]
assert(partition(A, 0, len(A)-1)==6)
```

## Question 3.

Code up your own Python implementation of `quicksort(A, p, r)`, using `partition(A,p,r)`.

<!-- grade_id: cell-8c39ebb8cd1aa83a, points: 0, grade: false, solution: true -->
```python
def quick_sort(A,p,r):
    if p >= r:
        return r
    else:
        q = partition(A, p, r) - 1
        quicksort(A, p, q)
        quicksort(A, q+2, r)
    return A
```

<!-- grade_id: cell-4f822430efd456ee, points: 1, grade: true, solution: false -->
```python
A = [0]
assert(quick_sort(A, 0, 0) == [0])
A = [3,1,2]
assert(quick_sort(A, 0, 2) == [1,2,3])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-6-3ff2a53494de> in <module>
      1 A = [0]
----> 2 assert(quick_sort(A, 0, 0) == [0])
      3 A = [3,1,2]
      4 assert(quick_sort(A, 0, 2) == [1,2,3])

AssertionError:
```

## Question 4. 
Explain (using experimental plots) the running time of `quick_sort` when: 
1. all elements of array A have the same value (e.g., [1,1,1])?
2. array A contains distinct elements sorted in decreasing order (e.g., [5,4,2,1])?


<!-- grade_id: cell-b58035dd5fa02329, points: 0, grade: true, solution: true -->
```python
# ATTEMPT 1

random.seed(123)

A = [1]*100
B = [i for i in range(0,100,-1)] # elements from 1-100 in decreasing order
start_A = time.time()
quick_A = quick_sort(A,0,len(A)-1)
end_A = time.time()
duration_A = end_A - start_A

start_B = time.time()
quick_B = quick_sort(B,0,len(B)-1)
end_B = time.time()
duration_B = end_B - start_B
```

```python
# ATTEMPT 2

def randomListGenerator(n):
    A=[]
    for i in range(n):
        A.append(random.randint(1,100)) # Generate a random number between 0 and 100
    return A 
```

```python
A_times = []
B_times = []

for i in range(1, 1001):
    A = randomListGenerator(i) 
    start_A = time.time()                  
    quick_sort(A,0,len(A)-1)
    end_A = time.time()
    duration_A = end_A - start_A              
    A_times.append(duration_A)

for i in range(1, 1001):
    B = randomListGenerator(i) 
    start_B = time.time()                  
    quick_sort(B,0,len(B)-1)
    end_B = time.time()
    duration_B = end_B - start_B              
    B_times.append(duration_B)
```

```python
x = np.linspace(1,1001,1000).astype(int)

plt.figure(figsize=(15,7.5))
plt.plot(x, A_times, color='red', label='A_times')    
plt.plot(x, B_times, color='blue', label='B_times')  
plt.title('Practical runtime comparison for different sorting algorithms')
plt.xlabel('Number of inputs')
plt.ylabel('Run time (s.)')
plt.legend()
plt.show()
```

```python

```