Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rebecca"
COLLABORATORS = "None"
```

---

# CS110 Pre-class Work 7.1

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson 3.2, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1.
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


```python
import heapq
import random

# Test array to check if I understand how heapq.heapify() works:
a = [1,2,3,6,7,89,2,3]
print("List a:", a)
b = [i * -1 for i in a] # List comprehension: trying to achieve the same thing as maxheapify, but unsure whether it's sorted...
print("List b:", b)

# NB: Do not rename the variable because heapify() works on the list in-place and will return NoneType
heapq.heapify(a)
print("List a as a heap:", a)
heapq.heappop(a)
print("Pop the root of a:", a)
heapq.heappush(a, 100)
print("Push a new element onto a:", a)

heapq.heapify(b)
print("List b as a heap:", b)
heapq.heappop(b)
print("Pop the root of b:", b)
heapq.heappush(b, 100)
print("Push a new element onto b:", b)
```

Output:
```
List a: [1, 2, 3, 6, 7, 89, 2, 3]
List b: [-1, -2, -3, -6, -7, -89, -2, -3]
List a as a heap: [1, 2, 2, 3, 7, 89, 3, 6]
Pop the root of a: [2, 2, 3, 3, 7, 89, 6]
Push a new element onto a: [2, 2, 3, 3, 7, 89, 6, 100]
List b as a heap: [-89, -7, -3, -6, -2, -1, -2, -3]
Pop the root of b: [-7, -6, -3, -3, -2, -1, -2]
Push a new element onto b: [-7, -6, -3, -3, -2, -1, -2, 100]
```

<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
# minh = upper side of the median
# maxh = lower side of the median

def add_to_median_heap(minh, maxh, elem):
    # Rememebr, the difference in length between the two heaps can never be greater than 1...
    if len(maxh) == 0:
        maxh.append(elem)
    # If maxh is longer:
    elif len(maxh) - len(minh) == 1:
        minh.append(elem)
        heapq.heapify(minh)
    # if minh is longer:
    elif len(minh) - len(maxh) == 1:
        maxh.append(elem)
        heapq._heapify_max(maxh)
    else:
        if elem <= maxh[0]:
            minh.append(elem)
            heapq.heapify(minh)
        else:
            maxh.append(elem)
            heapq._heapify_max(maxh)
    return(minh, maxh)
```

## Question 2
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    # Whichever heap has more elements must have the median in its root:
    if len(minh) > len(maxh):
        median = minh[0]
    elif len(maxh) > len(minh):
        median = maxh[0]
    else:
        median = (maxh[0] + minh[0])//2
    return median
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3.

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
# Clearly, something is wrong in my code:

minh = []
maxh = []
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1
2
5
4
9
6
13
8
17
10
21
12
25
14
29
16
33
18
37
20
41
22
45
24
49
26
53
28
57
30
61
32
65
34
69
36
73
38
77
40
81
42
85
44
89
46
93
48
97
50
```

## Question 4.
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
O(n)

## Question 5.
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
O(1) because we are only going through a series of conditions and accessing the first element in the list when one of these conditions is satisfied. The length of the lists will not affect this.

## Question 6.

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
I'm not quite sure what "vanilla" is referring to here, but I'll assume it applies to some variation of Quicksort. 

Using heaps is arguably more efficient for larger sizes of n. The most basic approach of first sorting the list via a pivot and then finding the middle element/s has a recurrence relationship of O(n log n) + O(1) [where the constant is for finding the median], therefore expected time complexity O(n log n) and a worst case time complexity of O(n^2). Assuming we are not dealing with randomized pivots, we can isolate our analysis to only the average case O(n log n) – because we know the worst case usually only occures when the pivot is the first or last element in the list, leaving virtually the entire list - 1 to be sorted in the same way again, hence resulting in n^2 time.

Using heaps is more efficient because you do not need to recursively sort the list, as long as the parent in minh is smaller and all its children and the parent in maxh is larger than all its children. This allows us to pop a single element a time to find the median.

## [Optional] Question 7.

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-7-15b94d1fa268> in <module>
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1.

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


```python
# ATTEMPT 1:

def qselect(lst, k): 
    pivot = [lst.pop()] # pops in-place # giving me problems in the test because it is "popping from an empty list"
    larger = [] # elements that are larger than the pivot value
    smaller = [] # elements that are smaller than the pivot value
    for num in lst:
        if num > pivot[0]:
            larger.append(num)
        elif num == pivot[0]:
            pivot.append(num)
        else:
            smaller.append(num)
    
    if len(smaller) < k <= len(smaller)+len(pivot):
        return pivot[0] # this is the kth smallest element
    elif len(smaller) >= k: # the pivot can't be the kth smallest element
        return qselect(smaller, k)
    else:
        return qselect(larger, k-len(smaller)-len(pivot))
```

<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
## ATTEMPT 2

def partition(A, start, end):
    pivot = end
    i = start-1
    for j in range(start, end):
        if A[j] <= A[pivot]:
            i += 1 # move on to next index
            A[i], A[j] = A[j], A[i] # swap the element yo just analyzed (j) with the next one
    A[i+1], A[pivot] = A[pivot], A[i+1] # new pivot is the element immediately after the one that was smaller than the first pivot
    return (i+1)

def qselect_2(A,start,end,k):
    if start >= end:
        return A[end]
    else:
        p = partition(A,start,end)
        print(A)
        if p == k-1:
            return A[p]
            print(A[p])
        elif p < k-1:
            i = p - 1
            print(A[i+2:end])
            qselect_2(A,i+2,end,k)          
        elif p > k-1:
            i = p - 1
            print(A[start:i])
            qselect_2(A,start,i,k)

import random
random.seed(123)
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect_2(lst1, 0, len(lst1)-1, a))
assert(lst2 == list(range(100)))
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

```python
# ATTEMPT 3 with random pivot

def select(lst, l, r, index):

    # base case
    if r == l:
        return lst[l]

    # choose random pivot
    pivot_index = random.randint(l, r)

    # move pivot to beginning of list
    lst[l], lst[pivot_index] = lst[pivot_index], lst[l]

    # partition
    i = l
    for j in xrange(l+1, r+1):
        if lst[j] < lst[l]:
            i += 1
            lst[i], lst[j] = lst[j], lst[i]

    # move pivot to correct location
    lst[i], lst[l] = lst[l], lst[i]

    # recursively partition one side only
    if index == i:
        return lst[i]
    elif index < i:
        return select(lst, l, i-1, index)
    else:
        return select(lst, i+1, r, index)

    if items is None or len(items) < 1:
        return None

    if item_index < 0 or item_index > len(items) - 1:
        raise IndexError()

    return select(items, 0, len(items) - 1, item_index)
```

## Question 2.
Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.


```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):
    print(qselect(lst, a))
```

## Question 3.

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
For Attempt #1:
Define pivot array = Big_theta(1)
For num in nlist = T(n)

If statements:
simple case = Big_theta(1)

Case 2 call on recurrence algorithm: T(n)

Case 3 call on recurrence algorithm: T(n)

Best case: Big_theta(1) + Big_theta(1) = Big_theta(1)

Worst case: T(n^2) + Big_theta(1)

## Question 4.

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
O(n) because the for loop goes through every element in the list regardless of the size of n.
len(smaller) < k <= len(smaller)+len(pivot) then we do not need to call on quickselect again, so the complexity is in linear time.

## Question 5.

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
O(n^2)
This accounts for the first for loop where you go through the entire list and one of the elif/else statement where you call qselect again.

```python

```