Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Rebecca"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work 5.1

## Question 1.

### Question 1a.

Write the worst sort function in the world, `worst_sort`. This function takes a list, and then randomly shuffles it until it happens to be in sorted order. Once it is in sorted order then the list is returned.


```python
import random
import math
import numpy as np
```

<!-- grade_id: cell-ead9c74054a1c5da, points: 0, grade: false, solution: true -->
```python
worst_step = [0]

def worst_sort(A):
    """
    Sort A in ascending order by randomly shuffling its 
    elements until they are in order.
    
    Input:
    - A: list of numerical values
    
    Output:
    - A: Sorted list
    """
    worst_step.append(1) # step for checking the while loop condition
    # As long as A is not sorted in either ascending or descending order, shuffle the list:
    while A != sorted(A) and A!= sorted(A).reverse():
        A = random.shuffle(A)
        worst_step.append(1)
    return A, len(worst_step)
```

```python
worst_sort([1,2,3])
worst_sort([1,2,3,4,5,6,10,4])
```

Output:
```
<ipython-input-10-e4276d6bc769> in <module>
      1 worst_sort([1,2,3])
----> 2 worst_sort([1,2,3,4,5,6,10,4])

<ipython-input-8-05cf3130bb91> in worst_sort(A)
     15     # As long as A is not sorted in either ascending or descending order, shuffle the list:
     16     worst_step.append(1)
---> 17     while A != sorted(A) and A!= sorted(A).reverse():
     18         A = random.shuffle(A)
     19         worst_step.append(1)

TypeError: 'NoneType' object is not iterable
```

<!-- grade_id: cell-ff4db3a3e8b04515, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 1b.
What is the best case complexity of this algorithm?

<!-- grade_id: cell-afac80171b3b6232, points: 0, grade: true, solution: true -->
The best case is where the first shuffle magically happens to sort the list in ascending order, which equals Ω(1).

### Question 1c.

What is the average case complexity?


<!-- grade_id: cell-9d52d34daa1e3478, points: 0, grade: true, solution: true -->
The number of permutations of a list is equal to the factorial of the length of the list: Θ(n!). This is the average case because we assume the pseudo random generator outputs results that have some kind of relationship to the input size.

### Question 1d.

For what size lists is this a feasible method?

<!-- grade_id: cell-ab40f08768579798, points: 0, grade: true, solution: true -->
This could be the reverse factorial of the largest number that a 32-bit Python console can handle, which is 2^31 - 1.

### [Optional] Question 1e.

Can you think of an even worse sorting method? In such a case, what would its complexity be? How big a list could you sort?

<!-- grade_id: cell-ac05806942341926, points: 0, grade: true, solution: true -->
[Not answering the question, but I am taking notes for my own understanding]

Does the worst_sort have a worst case? According to the readings, pseudo random number generators remove the risk of having a worst-case because nothing is dependent on input size. But this is the answer I had thought of, is it valid?


"Because we are using a pseudo random generator for each shuffle, the truth is we can't be certain that there is a relationship to input size n. Even a list such as [3,2,1] might go through 6,023 shuffles before it produces the sorted list [1,2,3]. This is because the random generator does not keep track of the combinations it has already covered. For this reason, the absolute worst case complexity is O(∞)."

## Question 2.
Approximate median finder. Given a list and δ (a number between 0 and 0.5), the approximate median finder function returns a number that is guaranteed to lie between the (50-δ/2)% and (50+δ/2)% percentiles. Implement such a function by randomly selecting an element from the list and testing whether or not it lies within the bounds. If it doesn’t then retry with a new random element, but only a limited number of retries are allowed (you decide the maximum number of retries.)

### Question 2a.
Complete the following function


<!-- grade_id: cell-2ab65512a6148d3c, points: 0, grade: false, solution: true -->
```python
def check_approx_med(A, num, delta):
    """
    Given a list and a number in a list, determine whether it is within (50-delta/2)% and 
    (50+delta/2)% percentiles.
    X% percentile of a list is defined to be the smallest number in the list that is as large as at least
    X% numbers in the list. 
    
    Input:
    - A: list of numerical values
    - num: the number of interest
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - isApproxMed: a boolean value indicating whether num is within the bound
    """

    # Define indices for the lower and upper percentiles:
    # (Note: Confused as to whether I'm returning a value or an index...)
    lower_bound = A(50 - delta/2)/100*len(A)
    upper_bound = A(round((50 + delta/2)/100*len(A))
    
    # Set boolean value:
    isApproxMed = False    
    
    if num < lower_bound or random_element > upper_bound: 
        isApproxMed = False
    else:
        isApproxMed = True
    
    while isApproxMed = False:
        # Try a new random element:
        num = random.choice(A)
    
    return isApproxMed
```

Output:
```
File "<ipython-input-47-63cb8ea8c9cc>", line 23
    isApproxMed = False
              ^
SyntaxError: invalid syntax
```

<!-- grade_id: cell-df56d7433abb6517, points: 1, grade: true, solution: false -->
```python
assert(check_approx_med([0,1], 0, .25) == True)
assert(check_approx_med([0], 0, .25) == False)
```

Output:
```
<ipython-input-48-aa1f56e94d4b> in <module>
      1 assert(check_approx_med([0,1], 0, .25) == True)
----> 2 assert(check_approx_med([0], 0, .25) == False)

<ipython-input-37-7ffe3c84639b> in check_approx_med(A, num, delta)
     19     # Define indices for the lower and upper percentiles:
     20     lower_bound = A[round((50 - delta/2)/100*len(A))]
---> 21     upper_bound = A[round((50 + delta/2)/100*len(A))]
     22 
     23     # Set counter which will keep track of how many tries we've had:

IndexError: list index out of range
```

### Question 2b.
Complete the following function that makes use of `check_approx_med` above.


<!-- grade_id: cell-e19685b545147dde, points: 0, grade: false, solution: true -->
```python
def approx_med_finder(A, delta):
    """
    Given a list, find a number in the list that is between 50-delta/2% and 50+delta/2% percentiles
    
    Input:
    - A: list of numerical values
    - delta: the error, lies between 0 and 0.5
    
    Output:
    - num: the approximated median, if it is found within 100//delta trials (this is one possible 
    maximum number of retries. While you are encouraged to play around with another limits, the 
    submitted work must use this number.)
    - None: finding failed, if nothing found within 100//delta trials

    Note: the 100//delta is chosen here as it is the expected number of trials for the first successful finding to occur (using geometric distribution)
    """
    num = random.choice(A)
    
    # Only delta//100 number of trials allowed
    for i in range(delta//100):
        answer = check_approx_med(A, num, delta)
        if answer == True:
            return answer
        else:
            print("None")
```

### Question 2c.

What is the probability of failure in each random trial? What is the probability of failure after all the allowed random trials? Does this scale with δ or N (the number of elements in a list)?


<!-- grade_id: cell-efe7ab76f6a2546a, points: 0, grade: true, solution: true -->
Probability of failure = 1 - probability of success:
1-(1/n!)

Probability of failure for all random trials:
Sum (1-(1/n!))

### Question 2d.
Analyze the expected runtime of `approx_med_finder`. Note that because the function uses `check_approx_med`, you will most likely need to analyze the runtime of that function, too.

<!-- grade_id: cell-86b199798477fc2a, points: 0, grade: true, solution: true -->
Because we use a random generator to choose an element:
Generate random number = O(1)
For loop in approx_med_finder = O(n)
Call on check_approx_med = O(n)

Total = O(n)*O(n) + O(1)

= O(n^2)

# Personal Notes

```python
# PERMUTE-BY-SORTING

def permute_by_sorting(A):
    n = len(A)
    P = [None]*n
    for i in range(n):
        P[i] = random.randrange(0, n^3)

        #
```

```python
# RANDOMIZE-IN-PLACE 
# QUESTION: How is this difference from "PERMUTE-WITH-ALL"?

def randomize_in_place(A):
    n = len(A)
    for i in range(n):
        exchange_index = A[randrange(i, n)]
        A[i], A[exchnage_index] = A[exchnage_index], A[i] 
    
```